### Arithmetic Expressions with Parsing Expression Grammar

Consider the following PEG for arithmetic expressions:

    S ← E
    E ← T ('+' T)*
    T ← F ('×' F)*
    F ← P'²' / P
    P ← 'a' / '(' E ')'

Implement it as a backtracking parser, following the rules from the course notes. You may use the commented version of method `literal`.

<!-- grade_id: cell-d26520c7eddbab9b, points: 3, grade: true, solution: true -->
```python
class Backtrack:
    src: str

    def literal(self, k: int, a: str):
        if self.src.startswith(a, k): return k + len(a)
        # if self.src.startswith(a, k): print('match', a, 'at', k); return k + len(a)
        # else: print('fail', a, 'at', k); return None

    # YOUR CODE HERE
    raise NotImplementedError()
    
    def parse(self, s: str):
        self.src = s; return self.S(0) == len(s)
```

```python
p = Backtrack()
```

<!-- grade_id: cell-2defb949e3f95c23, points: 1, grade: true, solution: false -->
```python
assert p.parse('a²')
assert p.parse('a²+a²')
assert p.parse('a²+a+a²')
```

<!-- grade_id: cell-100652ba17ebb90d, points: 1, grade: true, solution: false -->
```python
assert p.parse('((a)²)')
assert p.parse('(a²+a)×a')
assert p.parse('a×a²×a')
```

<!-- grade_id: cell-35c67e76c99ad61e, points: 1, grade: true, solution: false -->
```python
assert not p.parse('a²²')
assert not p.parse('a+²a')
assert not p.parse('a(+)a')
```