### Packrat Parsing for Statements

Consider the following PEG for assignment statements and procedure calls; it is a fragment of the P0 grammar:

    statement ← 
        ident selector ':=' ident /
        ident (',' ident)* ':=' ident (',' ident)* /
        ident (',' ident)* '←' ident '(' ident ')'
    selector ← ('[' ident ']' / '.' ident)*
    ident  ← 'a' / ... / 'z'

Implement the parser as a backtracking parser, following the rules from the course notes. You may use the commented version of method `literal`.

<!-- grade_id: cell-d26520c7eddbab9b, points: 3, grade: true, solution: true -->
```python
class StatementBacktrack:
    src: str

    def literal(self, k: int, a: str):
        if self.src.startswith(a, k): return k + len(a)
        # if self.src.startswith(a, k): print('match', a, 'at', k); return k + len(a)
        # else: print('fail', a, 'at', k); return None

    # YOUR CODE HERE
    raise NotImplementedError()
    
    def parse(self, s: str):
        self.src = s; return self.statement(0) == len(s)
```

```python
pb = StatementBacktrack()
```

Here are some tests:

<!-- grade_id: cell-2defb949e3f95c23, points: 1, grade: true, solution: false -->
```python
assert pb.parse('a:=b')
assert pb.parse('a.b.c:=d')
assert pb.parse('a[b].c[e]:=f')
assert pb.parse('a,b,c:=d,e,f')
assert pb.parse('a,b,c←d(e)')
assert not pb.parse('a')
assert not pb.parse('a:b')
assert not pb.parse('a.b.(c)')
assert not pb.parse('a..b:=c')
assert not pb.parse('a.b:=7')
```

Now write a memoizing parser:

<!-- grade_id: cell-a3251abc95ae165e, points: 1, grade: true, solution: true -->
```python
class StatementMemoizing(StatementBacktrack):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
pm = StatementMemoizing()
```

Here are the same tests:

<!-- grade_id: cell-784156b568a51952, points: 1, grade: true, solution: false -->
```python
assert pb.parse('a:=b')
assert pb.parse('a.b.c:=d')
assert pb.parse('a[b].c[e]:=f')
assert pb.parse('a,b,c:=d,e,f')
assert pb.parse('a,b,c←d(e)')
assert not pb.parse('a')
assert not pb.parse('a:b')
assert not pb.parse('a.b.(c)')
assert not pb.parse('a..b:=c')
assert not pb.parse('a.b:=7')
```